Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

NAME = "Tanaka"
COLLABORATORS = ""

---

---
## Lab 8: Git



---
### Exercise 1: Git Together
---

![merge_conflict.png](attachment:0a481736-1382-4983-bda3-bba69da54a01.png)

Git merge is an important feature which allows you to integrate changes from one branch into another, allowing developers to work on different branches and then merge later. Merge conflicts can occur due to various reasons. A merge conflict arises when two separate branches have made edits to the same line in a file, or when a file has been deleted in one branch but edited in the other. In this exercise, we only consider this case: Two developers attempt to change the same line in a file and trigger a merge conflict when they commit. For simplicity, the two developers will not add or delete lines, but only make changes on existing lines of the file.

Write a function merge() that takes three files as parameters: the first is the original file, the second is the commit made by the first developer, and the third file is the commit made by the second developer. If the same line in the original file is changed by two developers, it is considered a conflict. Your function should return a boolean variable describing whether there will be a conflict or not. See the test cases below for examples. The files are available on Moodle in the “Lab” tile under Lab 8.

merge(“original.txt”, “noconflict_a.txt”, “noconflict_b.txt”) returns False

merge(“original.txt”, “noconflict_c.txt”, “noconflict_d.txt”) returns False

merge(“original.txt”, “conflict_a.txt”, “conflict_b.txt”) returns True

<!-- grade_id: Part1, points: 0, grade: false, solution: true -->
```python
def merge(original_file, file_2, file_3):
    original_file_1 = open(original_file, "r")
    original_file_c = original_file_1.readlines()
    original_file_1.close()

    file_2_open = open(file_2, "r")
    file_2_c = file_2_open.readlines()
    file_2_open.close()

    file_3_open = open(file_3, "r")
    file_3_c = file_3_open.readlines()
    file_3_open.close()

   
    for line in range(len(original_file_c)):
        if file_2_c[line] != original_file_c[line] and file_3_c[line] != original_file_c[line]:
            if file_2_c[line] != file_3_c[line]:
                return True 

    return False

print(merge("original.txt", "noconflict_a.txt", "noconflict_b.txt"))  
print(merge("original.txt", "noconflict_c.txt", "noconflict_d.txt")) 
print(merge("original.txt", "conflict_a.txt", "conflict_b.txt"))     












```

Output:
```
False
False
True
```

---
### Exercise 2: Git or Miss
---

The Git commands required to push or pull code to or from a remote repository can vary based on factors such as your current working directory, active branches, and the state of the repository. For this exercise, we will work with a fixed set of commands that must be followed in a specific order to perform either a push or a pull.

**In order to push code we assume the following 5 commands must be executed in order:**

git init -b [branch_name] # initialize the local directory as a Git repository

git add [file_name] # add the files in your new local repository

git commit -m [message] # Commits the tracked changes and prepares them to be pushed to a remote repository

git remote add origin [remote_url] # add the URL for the remote repository where your local repository will be pushed

git push -u origin [branch_name] # push the changes in your local repository to remote


**To pull code, one of the following sets of commands must be executed in order:**

**1st set:**

git checkout [branch_name] # checkout and switch to the branch

git pull [remote_repo] # pull down the branch

**2nd set:**

git checkout main # switch to the main branch

git pull --rebase origin # rebase and synchronize with the remote main branch.


Write a function git_action() that reads a series of Git commands from a text file and determines whether these commands, when executed sequentially, result in a push, pull, or neither. Your function should validate the commands and their flags (excluding the parts in square brackets, we assume these parts must exist, but can be any string other than empty) and return "push", "pull", or "neither" as a string. The function should take the filename as input.

**Sample Test Cases (get files from Moodle):**

git_action(“push.txt”) returns push

git_action(“pull_a.txt”) returns pull

git_action(“pull_b.txt”) returns pull

git_action(“neither_a.txt”) returns neither

git_action(“neither_b.txt”) returns neither

git_action(“neither_c.txt”) returns neither

<!-- grade_id: Part2, points: 0, grade: false, solution: true -->
```python
def git_action(file_name):
    file = open(file_name, "r")
    file_r = file.readlines()
    file.close()
    
    push_sequence = [
        "git init -b",
        "git add",
        "git commit -m",
        "git remote add origin",
        "git push -u origin"
    ]
    
    pull_sequence_a = [
        "git checkout",
        "git pull"
    ]
    
    pull_sequence_b = [
        "git checkout main",
        "git pull --rebase origin"
    ]
    
    
    commands = [" ".join(line.split()[:3 if "git init -b" in line else 2]) for line in file_r]
    
    if commands == push_sequence:
        return "push"
    elif commands == pull_sequence_a or commands == pull_sequence_b:
        return "pull"
    else:
        return "neither"
git_action("push.txt")
git_action("pull_b.txt")
 
```

Output:
```
'pull'
```

---
### Exercise 3: The Commands
---

Read the "Git Cheat Sheet" under Week 10 Course Materials on Moodle, and answer the following questions. You can simply write the command down as text.

#### 3.1. What is the command to initialize a git repository?

<!-- grade_id: Part3-1, points: 2, grade: true, solution: true -->
git init

#### 3.2 What is the command to show modified files in working directory?

<!-- grade_id: Part3-2, points: 2, grade: true, solution: true -->
git status

#### 3.3 What is the command to add a file called "file.txt" to the staging environment?

<!-- grade_id: Part3-3, points: 2, grade: true, solution: true -->
git add 

#### 3.4 What is the command to create a new branch called "new-branch"?

<!-- grade_id: Part3-4, points: 2, grade: true, solution: true -->
git branch [branch-name]

#### 3.5 What is the command to push local commits on "new-branch" to a remote repository?

<!-- grade_id: Part3-5, points: 2, grade: true, solution: true -->
git push origin

#### 3.6 What is the command to get changes of "new-branch" on a remote repository back to local machine?

<!-- grade_id: Part3-6, points: 2, grade: true, solution: true -->
git fetch origin new-branch
git checkout new-branch

---
### Exercise 4: Git It Out There
---

Create a repository on Github called "CS128Lab8".

Commit this notebook to that repository.

There are multiple ways to do that. One way is to commit straight from the command line in Jupyter. Another (maybe) simpler way is to download this notebook, go to your github repo, and upload through "add files".

After your notebook is put on Github, paste the link to your repository below:


<!-- grade_id: Part4, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-9e5a69bc78d1d35e, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(merge("original.txt", "noconflict_a.txt", "noconflict_b.txt"), False)

### END TEST
```

<!-- grade_id: cell-e3d94519613af313, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(merge("original.txt", "noconflict_c.txt", "noconflict_d.txt"), False)

### END TEST
```

<!-- grade_id: cell-ac3dd8fd1b2633f1, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(merge("original.txt", "conflict_a.txt", "conflict_b.txt"), True)

### END TEST
```

<!-- grade_id: cell-e1bd57119a701315, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("push.txt"), "push")

### END TEST
```

<!-- grade_id: cell-d92798150c85943b, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("pull_a.txt"), "pull")

### END TEST
```

<!-- grade_id: cell-4b0304d56fbd5c55, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("pull_b.txt"), "pull")

### END TEST
```

<!-- grade_id: cell-879f8a35b910086f, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("neither_a.txt"), "neither")

### END TEST
```

<!-- grade_id: cell-6486f94ce5153fe2, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("neither_b.txt"), "neither")

### END TEST
```

<!-- grade_id: cell-6b141a2e76616274, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(git_action("neither_c.txt"), "neither")

### END TEST
```